##### CAUTION: Completing and submitting the assignment constitutes your commitment to complete the assignment independently and not to plagiarize or provide others with the opportunity to plagiarize.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "郭子涵"
COLLABORATORS = "0921201"
```

---

### Problem 1:Create Maximum Number (40 points)
Given two arrays of length m and n with digits 0-9 representing two numbers. Create the maximum number of length k <= m + n from digits of the two. The relative order of the digits from the same array must be preserved. Return  a number consisting of k digits.

Example：  
Input:  
nums1 = [3, 4, 6, 5]  
nums2 = [9, 1, 2, 5, 8, 3]  
k = 5  
Output:  
98653

<!-- grade_id: cell-312688a68d66ff1f, points: 0, grade: false, solution: true -->
```python
def max_number(nums1,nums2,k):
    n = len(nums1)
    m = len(nums2)
    # YOUR CODE HERE
    def pick_max(nums, k):
        stack = []
        rem = len(nums) - k
        for num in nums:
            while rem and stack and stack[-1] < num:
                stack.pop()
                rem -= 1
            stack.append(num)
        return stack

    def merge(nums1, nums2):
        ans = []
        while nums1 or nums2:
            bigger = nums1 if nums1 > nums2 else nums2
            ans.append(bigger[0])
            bigger.pop(0)
        return ans

    len1, len2 = len(nums1), len(nums2)
    ans = []
    for i in range(max(0, k - len2), min(k, len1) + 1):
        tmp = merge(pick_max(nums1, i), pick_max(nums2, k - i))
        ans = max(ans, tmp)
    return int(''.join(map(str, ans)))
```

<!-- grade_id: cell-2f586e928536e9c3, points: 40, grade: true, solution: false -->
```python
# max_number function test
nums1 = [3, 4, 6, 5]
nums2 = [9, 1, 2, 5, 8, 3]
k = 5
max_num = max_number(nums1,nums2,k)
assert max_num == 98653

```

### Problem 2:LIS (60 points)
Given an unsorted array of integers, find a length of longest increasing subsequence.

Example:  
Input: [10,9,2,5,3,7,101,18]  
Output: [2,3,7,101] (or other 4 items increasing subsequence)  
Explanation: The longest increasing subsequence is [2,3,7,101], therefore the length is 4.  
Please design and explain your Dynamic Programming recursive formula, then implement the algorithms in Python.

Caution: Use LCS Algorithm to Solve This Problem is NOT Allowed!

<!-- grade_id: cell-2ccc8c16e2e45724, points: 0, grade: false, solution: true -->
```python
# find the longest increasing subsequence
def lis(A):
    # YOUR CODE HERE
    if not A:
        return []
    n = len(A)
    dp = [1] * n
    for i in range(1, n):
        for j in range(i):
            if A[i] > A[j]:
                dp[i] = max(dp[i], dp[j] + 1)
    max_length = max(dp)
    res = []
    for i in range(n - 1, -1, -1):
        if dp[i] == max_length:
            res.append(A[i])
            max_length -= 1
    return res[::-1]
```

<!-- grade_id: cell-fadd57aea094cf93, points: 60, grade: true, solution: false -->
```python
# lis function test
# The return result of the function should be the longest increasing subsequence array。
A = [10,9,2,5,3,7,101,18]
print(lis(A))
assert(len(lis(A)) == 4)
```

Output:
```
[2, 3, 7, 18]
```