##### CAUTION: Completing and submitting the assignment constitutes your commitment to complete the assignment independently and not to plagiarize or provide others with the opportunity to plagiarize.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "郭子涵"
COLLABORATORS = "0921201"
```

---

### Problem 1: 将k个有序序列合并为整体有序 (30 points)
 (1) 考虑这样一个对k个有序序列（每个序列中含n/k个元素）进行合并排序的算法，将第1个序列与第2个进行线性时间内的Merge操作（就像我们在合并排序中使用的那样。然后，将合并得到的序列（含2n/k个元素）与第3个有序序列合并，接着再用得到的含3n/k个元素的序列与第4个有序序列合并... 直到最终得到一个含有所有元素的有序序列，请实现该算法，并分析其时间复杂性；
 
 (2) 请用分治法解决该问题并分析其时间复杂性。

首先完成merge函数。 (10 points)

<!-- grade_id: cell-d56db2b6577e42ce, points: 0, grade: false, solution: true -->
```python
def merge(A,p,q,r):
    s=p
    t=q+1
    k=p
    B=[0]*len(A)

    while s<=q and t<=r:
    # YOUR CODE HERE
        if A[s]<=A[t]:
            B[k]=A[s]
            s+=1
        else:
            B[k]=A[t]
            t+=1
        k+=1
    if s==q+1:
        B[k:r+1]=A[t:r+1]
    else:
        B[k:r+1]=A[s:q+1]
    A[p:r+1]=B[p:r+1]
```

<!-- grade_id: cell-2f586e928536e9c3, points: 10, grade: true, solution: false -->
```python
# 测试 merge函数
A = [2,4,6,1,3,5]
merge(A,0,2,5)
assert A == [1,2,3,4,5,6] 

```

设计k_merge函数，将k段有序序列（每个序列中含n/k个元素）合并为整体有序。 (8 points)

<!-- grade_id: cell-80e1684c1bc93854, points: 0, grade: false, solution: true -->
```python
def k_merge(A,k):
    # YOUR CODE HERE
    n = len(A)
    size = n // k
    for i in range(k-1):
        merge(A,i*size,(i+1)*size-1,(i+2)*size-1)
    merge(A,(k-1)*size,n-1,n-1)
    return A


```

<!-- grade_id: cell-fadd57aea094cf93, points: 8, grade: true, solution: false -->
```python
# 测试 k_merge函数
A = [1,3,2,6,4,8,5,7]
k_merge(A,4)
assert A == [1,2,3,4,5,6,7,8] 

```

k_merge的时间复杂性应为： (2 points)

(a)  $\Theta (log\mathit{n} )$

(b)  $\Theta (\mathit{n}* \mathit{k})$

(c)  $\Theta (\mathit{n}log\mathit{k} )$

(d)  $\Theta (\mathit{n}^2* \mathit{k} )$

##### 请根据你的选择，在下面的代码框中为 T_n 赋值 'a', 'b', 'c', 或 'd' ，例如：
T_n = 'd'

<!-- grade_id: cell-9230bbefcb0c3f8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
T_n = ''
```

<!-- grade_id: cell-2b9130acd2489101, points: 2, grade: true, solution: false -->
```python
# This test checks that your answer has the right format
# It does not check for correctness
assert T_n in ['a','b','c','d']

```

请用分治法解决该问题。（8 points）

<!-- grade_id: cell-df5d59afbb7fb452, points: 0, grade: false, solution: true -->
```python
def mySort_DC(A,low,high,k):
    # YOUR CODE HERE
    if high-low+1<=k:
        merge(A,low,(low+high)//2,high)
    else:
        mySort_DC(A,low,low+k-1,k)
        mySort_DC(A,low+k,high,k)
        k_merge(A,k)
    return A

```

<!-- grade_id: cell-364a37d522c63f48, points: 8, grade: true, solution: false -->
```python
# 测试 mySort_DC函数
A = [1,3,2,6,4,8,5,7]
mySort_DC(A,0,len(A)-1,4)
assert A == [1,2,3,4,5,6,7,8] 

```

mySort_DC的时间复杂性应为： (2 points)

(a)  $\Theta (log\mathit{n} )$

(b)  $\Theta (\mathit{n}* \mathit{k})$

(c)  $\Theta (\mathit{n}*log\mathit{k} )$

(d)  $\Theta (\mathit{n}^2* \mathit{k} )$

##### 请根据你的选择，在下面的代码框中为 T_n 赋值 'a', 'b', 'c', 或 'd' ，例如：
T_n = 'd'

<!-- grade_id: cell-9fbaaf8d614a8eb0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
T_n = ''
```

<!-- grade_id: cell-bf726c598319a162, points: 2, grade: true, solution: false -->
```python
# This test checks that your answer has the right format
# It does not check for correctness
assert T_n in ['a','b','c','d']

```

### Problem 2: 寻找多数元素 (30 points)
 (1) 给定一个大小为n的数组，请设计一个分治算法找到其中的多数元素，并分析其时间复杂性。多数元素是指在数组中出现次数大于⌊n/2⌋的元素，假设数组非空，且总是存在多数元素。

 （注：容易想到的分治策略是将元素对半划分，递归调用得到左、右半段的多数元素，整体的多数元素必为其一。但您可以想想，是否有更佳的分治策略？）

  (2) 请接着用哈希表完成该问题的算法设计，并分析其时间复杂性。

首先根据容易想到的分治策略设计majorityElement1函数。（8 points）

<!-- grade_id: cell-82efe0dae425e540, points: 0, grade: false, solution: true -->
```python
def majorityElement1(A,low,high):
    if low==high:
        return A[low]
    mid=(low+high)//2
    left=majorityElement1(A,low,mid)
    right=majorityElement1(A,mid+1 ,high)
    # YOUR CODE HERE
    if left==right:
        return left
    else:
        left_count=0
        right_count=0
        for i in range(low,high+1):
            if A[i]==left:
                left_count+=1
            elif A[i]==right:
                right_count+=1
        if left_count>right_count:
            return left
        else:
            return right
```

<!-- grade_id: cell-4a0f2e2b444392b9, points: 8, grade: true, solution: false -->
```python
# 测试 majorityElement1函数
A = [2,2,1,1,1,2,2]
m=majorityElement1(A,0,len(A)-1)
assert m == 2 

```

majorityElement1的时间复杂性应为： (2 points)

(a)  $\mathit{O} (log\mathit{n} )$

(b)  $\mathit{O} (\mathit{n} )$

(c)  $\mathit{O} (\mathit{n}log\mathit{n} )$

(d)  $\mathit{O} (\mathit{n}^2 )$

##### 请根据你的选择，在下面的代码框中为 T_n 赋值 'a', 'b', 'c', 或 'd' ，例如：
T_n = 'd'

<!-- grade_id: cell-6e48250c277166a1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
T_n = ''
```

<!-- grade_id: cell-2dd1592352892d1e, points: 2, grade: true, solution: false -->
```python
# This test checks that your answer has the right format
# It does not check for correctness
assert T_n in ['a','b','c','d']

```

接下来请思考是否有更佳的分治策略，并设计majorityElement2函数。（10 points）

<!-- grade_id: cell-42e4154e0cf67a50, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

def majorityElement2(A,low ,high):
    # YOUR CODE HERE
    if low==high:
        return A[low]
    mid=(low+high)//2
    left=majorityElement2(A,low,mid)
    right=majorityElement2(A,mid+1 ,high)
    if left==right:
        return left
    else:
        left_count=0
        right_count=0
        for i in range(low,high+1):
            if A[i]==left:
                left_count+=1
            elif A[i]==right:
                right_count+=1
        if left_count>right_count:
            return left
        else:
            return right
```

<!-- grade_id: cell-6c190a67c94e8dd5, points: 10, grade: true, solution: false -->
```python
# 测试 majorityElement2函数
A = [2,2,1,1,1,2,2]
m=majorityElement2(A,0,len(A)-1)
assert m == 2 

```

请用哈希表完成该问题的算法设计。（8 points）

<!-- grade_id: cell-5c2708e24aa39006, points: 0, grade: false, solution: true -->
```python
def majorityElement3(A,low ,high):
    Y={}
    for i in range(len(A)):
        # YOUR CODE HERE
        if A[i] in Y:
            Y[A[i]]+=1
        else:
            Y[A[i]]=1
    for i in Y:
        if Y[i]>len(A)/2:
            return i
    return None

```

<!-- grade_id: cell-699150e2e5b55d05, points: 8, grade: true, solution: false -->
```python
# 测试 majorityElement3函数
A = [2,2,1,1,1,2,2]
m=majorityElement3(A,0,len(A)-1)
assert m == 2 

```

majorityElement3的时间复杂性应为： (2 points)

(a)  $\Theta (log\mathit{n} )$

(b)  $\Theta (\mathit{n} )$

(c)  $\Theta (\mathit{n}log\mathit{n} )$

(d)  $\Theta (\mathit{n}^2 )$

##### 请根据你的选择，在下面的代码框中为 T_n 赋值 'a', 'b', 'c', 或 'd' ，例如：
T_n = 'd'

<!-- grade_id: cell-40476b465edcf405, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
T_n = ''
```

<!-- grade_id: cell-9d02273f08a3ea6e, points: 2, grade: true, solution: false -->
```python
# This test checks that your answer has the right format
# It does not check for correctness
assert T_n in ['a','b','c','d']

```

### Problem 3: 最佳小组问题 (40 points)
假设有n个职员，序号为0至n-1，其能力值记录在数组capacities中，现需从n个职员中选出一个最佳小组，要求小组中的职员序号必须连续，且他们的能力值之和在所有可能的小组中最大（小组规模最少1人，最多n人）。由于职员的能力值可为负，因此并非规模越大越好。请设计一个分治算法求出最佳小组所能得到的那个最大能力值，并分析其时间复杂性。

示例：

输入：capacities= [-2,1,-3,4,-2,3,-1,-5,4]

输出：5

因为取连续职员 [4,-2,3]形成的小组可得到总能力5。

请设计一个分治算法求出最佳小组所能得到的最大能力值。(35 points)

<!-- grade_id: cell-cd7497bf3ceefaaa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


def maxSubArray(A,low, high) :
    # YOUR CODE HERE
    if low==high:
        return A[low]
    mid=(low+high)//2
    left=maxSubArray(A,low,mid)
    right=maxSubArray(A,mid+1,high)
    mid_left=A[mid]
    mid_right=A[mid+1]
    temp=0
    for i in range(mid,low-1,-1):
        temp+=A[i]
        if temp>mid_left:
            mid_left=temp
    temp=0
    for i in range(mid+1,high+1):
        temp+=A[i]
        if temp>mid_right:
            mid_right=temp
    mid_max=mid_left+mid_right
    return max(left,right,mid_max)
```

<!-- grade_id: cell-f078c84263964963, points: 35, grade: true, solution: false -->
```python
# 测试 maxSubArray函数
A = [-2,1,-3,4,-2,3,-1,-5,4]
m=maxSubArray(A,0,len(A)-1)
assert m == 5 

```

maxSubArray的时间复杂性应为： (5 points)

(a)  $\Theta (log\mathit{n} )$

(b)  $\Theta (\mathit{n} )$

(c)  $\Theta (\mathit{n}log\mathit{n} )$

(d)  $\Theta (\mathit{n}^2 )$

##### 请根据你的选择，在下面的代码框中为 T_n 赋值 'a', 'b', 'c', 或 'd' ，例如：
T_n = 'd'

<!-- grade_id: cell-d1805d4d416844e2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
T_n = ''
```

<!-- grade_id: cell-050ea18bc6c3d329, points: 5, grade: true, solution: false -->
```python
# This test checks that your answer has the right format
# It does not check for correctness
assert T_n in ['a','b','c','d']

```